```python
# Initialize Otter
import otter
grader = otter.Notebook("t01.ipynb")
```

<img src = 'https://raw.githubusercontent.com/Delarbol/semillero/master/Banner%20DS.png'>
<h1 align=center style="font-size:3vw">Introducción a la Programación con Python</h1>
<h2 align=center>Educación Continua</h2>

---

# Ejercicio 1

Usted es un climatólogo en Colombia y quiere conocer los históricos de temperatura de los Estados Unidos. Al revisar las bases de datos se da cuenta que las temperaturas están dada en grados **Fahrenheit**. Por lo tanto, es necesario escribir un programa que al ingresar la temperatura en grados Farenheit la convierta a grados centígrados.

Cuando entregue su programa **es necesario** que use los siguientes valores:

**INPUT ->** temperatura (F°): 96.8  


**OUTPUT ->** Temp en C°: 36  -----> Este es el resultado deseado

### ¡NO CAMBIE EL NOMBRE DE LAS VARIABLES!
*No use inputs asigne un valor a fahrenheit ej. fahrenheit = 96.8

```python
fahrenheit= ...
print(celsius)
```

```python
grader.check("q_1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(filtering=False, pdf=False, run_tests=True)
```

Tests for q_1 (points: None)
```
>>> assert type(fahrenheit) == float
```
failure_message: fahrenheit deberia ser un float
success_message: Buen Trabajo
```
>>> assert celsius == 36.0
>>> 
>>> 
>>>
```
failure_message: hay algo mal en tu formula de conversion de grados revisala
success_message: Buen Trabajo
