```python
# Initialize Otter
import otter
grader = otter.Notebook("Tarea 1 - Variables.ipynb")
```

<img src = 'https://raw.githubusercontent.com/Delarbol/semillero/master/Banner%20DS.png'>
<h1 align=center style="font-size:3vw">Introducción a la Programación con Python</h1>
<h2 align=center>Educación Continua</h2>

---

# Ejercicio 1

Usted es un climatólogo en Colombia y quiere conocer los históricos de temperatura de los Estados Unidos. Al revisar las bases de datos se da cuenta que las temperaturas están dada en grados **Fahrenheit**. Por lo tanto, es necesario escribir un programa que le solicita al usuario ingresar la temperatura en grados Farenheit y la convierta a grados centígrados.

Cuando entregue su programa **es necesario** que use los siguientes valores:

**INPUT ->** Inserte temperatura (F°): 96.8  


**OUTPUT ->** Temp en C°: 36  -----> Este es el resultado deseado

### ¡NO CAMBIE EL NOMBRE DE LAS VARIABLES!

*Si la variable se llama "fahrenheit",no altere ese nombre

```python
fahrenheit= ...
print(celsius)
```

```python
grader.check("q_1")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(filtering=False, pdf=False, run_tests=True)
```

Tests for q_1 (points: 1)
```
>>> type(fahrenheit) == float
True
```
failure_message: deberías devolver tu respuesta como un float
success_message: ¡Buen Trabajo!
```
>>> celsius == 36.0
>>> 
>>> 
>>> 
True
```
failure_message: hay algo mal en tu formula de conversión de grados revísala
success_message: ¡Buen Trabajo!
